For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

<!-- grade_id: pedal_squares, points: 1, grade: true, solution: false -->
```python
from pedal import *
from pedal.environments.nbgrader import setup_environment

if setup_environment('squares'):
    prevent_ast("If")
    prevent_ast("If")
    prevent_embedded_answer(1)
    prevent_import('numpy')

```

Output:
```
<IPython.core.display.Javascript object>
c:\users\acbart\projects\pedal-edu\pedal\pedal\environments\nbgrader.py in force_immediate_errors(self, feedback, report)
    208                 exc_type, value, traceback = sys.exc_info()
    209                 traceback.tb_next = None
--> 210                 raise exc_type(value).with_traceback(traceback) from None
    211 
    212         #if in_notebook():

c:\users\acbart\projects\pedal-edu\pedal\pedal\environments\nbgrader.py in force_immediate_errors(self, feedback, report)
    203         with tc.subTest():
    204             try:
--> 205                 tc.fail(final.message)
    206             except:
    207                 # Remove traceback info as we don't need it

AssertionError: You used an IF statement on line 7. You may not use that.
```

```python

```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```